---
title: example_r_challenges
description: >-
  undefined


---
## True basics

```yaml
type: MultipleChoiceChallenge

xp: NaN

key: 167A4VcDii9Ku
```














---
## Identify the class of an object

```yaml
type: OutputChallenge

xp: NaN

key: 167oSz8A19yBt
```














---
## Create a vector

```yaml
type: BlanksChallenge

xp: NaN

key: 167y3A7MS2X70
```














---
## test 1

```yaml
type: NormalExercise

xp: NaN

key: 83259338a7
```

you have stuff to do

`@instructions`
here be instructions

`@hint`
no hints !!

`@pre_exercise_code`
```{}
# try to load the forcast package
library(forecast)
```
`@sample_code`
```{}
# no sample code provided
```
`@solution`
```{}
x <- 100
```
`@sct`
```{}
x = 100
```




