---
title: example_r_challenges
description: >-
  undefined


---
## True basics

```yaml
type: MultipleChoiceChallenge

xp: undefined

key: 167A4VcDii9Ku
```














---
## Identify the class of an object

```yaml
type: OutputChallenge

xp: undefined

key: 167oSz8A19yBt
```














---
## Create a vector

```yaml
type: BlanksChallenge

xp: undefined

key: 167y3A7MS2X70
```














---
## Test 1

```yaml
type: NormalExercise

xp: 

key: 238fab4426
```

some assignment text here

`@instructions`
generate a single random number in the range (0,1) and assign the result to the value x

`@hint`
no hints !!

`@pre_exercise_code`
```{}
'Woohoo'
```
`@sample_code`
```{}
# and no sample code
```
`@solution`
```{}
x <- runif(1)
```
`@sct`
```{}
x < 1
x > 0

```




