---
title: example_r_challenges
description: >-
  undefined


---
## True basics

```yaml
type: MultipleChoiceChallenge

xp: NaN

key: 167A4VcDii9Ku
```














---
## Identify the class of an object

```yaml
type: OutputChallenge

xp: NaN

key: 167oSz8A19yBt
```














---
## Create a vector

```yaml
type: BlanksChallenge

xp: NaN

key: 167y3A7MS2X70
```














---
## Test 1

```yaml
type: NormalExercise

xp: NaN

key: 238fab4426
```

some assignment text here

`@instructions`
generate a single random number in the range (0,1) and assign the result to the value x

`@hint`
no hints !!

`@pre_exercise_code`
```{undefined}
'Woohoo'
```
`@sample_code`
```{undefined}
# and no sample code
```
`@solution`
```{undefined}
x <- runif(1)
```
`@sct`
```{undefined}
x < 1
x > 0
```





---
## test 1

```yaml
type: NormalExercise

xp: 

key: 83259338a7
```

you have stuff to do

`@instructions`
here be instructions

`@hint`
no hints !!

`@pre_exercise_code`
```{}
# try to load the forcast package
library(forecast)
```
`@sample_code`
```{}
# no sample code provided
```
`@solution`
```{}
 x <- 100
```
`@sct`
```{}
x = 100
```




